# variables1.py
# Make me compile! Execute the command `pythonlings hint variables1` if you want a hint :)

# About this `I AM NOT DONE` thing:
# We sometimes encourage you to keep trying things on a given exercise,
# even after you already figured it out. If you got everything working and
# feel ready for the next exercise, remove the `I AM NOT DONE` comment below.

# I AM NOT DONE

# make sure x has a value 5
x = 5
print("x has a value {}".format(x))